```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 8 - Functions and Control

## Data 6, Fall 2024

Welcome to lab 8! This week we will be covering boolean operators and functions.

```python
# Just run this cell
from datascience import *

import warnings
warnings.simplefilter('ignore')
```



# Boolean Operators

---

## Comparison Operators

Comparison operators are boolean operators that are used to compare two different values in Python.

We introduced six of these operators in lecture: `==`, `!=`,`<`, `<=`, `>`, and `>=`. Try using these operators on variables and values of all types, though in practice they are typically used with integers and floats, and sometimes on strings. Strings are treated in increasing alphabetical order: ('a' < 'b') == True.

Run the following series of cells to see some examples of these comparison operators in practice:

```python
1 < 2
```

```python
min(20, 30) <= max(-10, 20)
```

```python
(3 + 4) != (3 * 4)
```

```python
"abc" < "def"
```

### An Important Distinction

There is an important distinction here between `=` and `==`:

>`=` is used to **set** a Python name (i.e. assignment), while `==` is used as the **comparison operator** to compare two values. 

```python
# Set the name 'a' to have value 5
a = 5
```

```python
# Ask if it is equal to the value you think it should be
a == 5
```

## Membership Operators: `in`
There is one more boolean operator built into Python we have not mentioned yet. The keyword `in` allows you to check if strings are contained in larger strings. It can be used in other ways in Python, but for now we will use it to ask if the first string is a substring of the second string (if it appears in the other string).

Run the cells below for examples of `in`:

```python
"hello" in "hello world"
```

```python
"i" in "rhythm"
```

The `in` keyword can also check whether an array *contains* a certain element. `True` is returned if the array contains the element, and `False` otherwise.

```python
my_array = make_array(3,2,5,6)
5 in my_array
```

```python
9 in my_array
```

## Boolean Operators: `and`, `or`, and `not`

`and`, `or`, and `not` are boolean operators that operate directly on booleans.

>`and` returns `True` when both operands are `True`, `False` otherwise.

>`or` returns `False` when both operands are `False`, `True` otherwise.

>`not` gives back the opposite boolean.

These three operators are useful when writing functions because they help control a function's logic. You will find them useful when writing more complicated functions in the course.

**Tip**: use parentheses to ensure your boolean operators are being applied correctly and without ambiguity.

Run the cells below to see these 3 operators in practice:

```python
(3 * 4 == 12) or (3 + 4 == 50)
```

```python
(3 * 4 == 12) and (3 + 4 == 50)
```

```python
not (3 + 4 == 50)
```

---
### Question 1

Write code that assigns `x_satisfies` to `True` or `False` depending on if the value of `x` satisfies *either* of the following conditions:
1. Even and has a value in the range 24 and 32 (inclusive)
2. Odd and has a value between 7 and 20 (exclusive). 

For this question, `x` will be assigned to `26`.

*Hint*: You should try using the modulo operator (%) to help check if a value is odd or even! The `%` operator will provide you with the *remainder* after diving by a certain number


```python
x = 26
x_satifies = ...
x_satifies
```

```python
grader.check("q1")
```

---
### Should I Stay or Should I Go
Now let's take a look at a function that uses **variables** and **boolean operators**:

Based on the color of the stoplight in front of us we want to know whether we should stop or keep going. Because we know we are going to encounter many lights on our trip, we should write a function! We can make two functions: one where we take risks and go on yellow lights, and one where we play it safe and stop for yellow lights

```python
def stop_or_go_risky(light):
    green = (light == 'green')
    yellow = (light == 'yellow')
    red = (light == 'red')
    return (green or yellow) and (not red) # Go on green and yellow, stop on red

def stop_or_go_safe(light):
    green = (light == 'green')
    yellow = (light == 'yellow')
    red = (light == 'red')
    return green and ((not yellow) and (not red)) # Go on green, stop on yellow and red
```

These functions are booleans because they output either True or False. We can write functions that output any type we need, and you will do so throughout the course!

Run the following cell as a demo:

```python
#We can use the outputs of these functions to let us know to stop the car if we need to on our way to the market.

light1 = stop_or_go_safe('green')
light2 = stop_or_go_safe('yellow')
light3 = stop_or_go_safe('red')

print("I am driving to the market! No rush!")
print("") # This is the same as hitting the return key while typing in Google Docs
print("Should I go at a green light?... " + str(light1))
print("Should I go at a yellow light?... " + str(light2))
print("Should I go at a red light?... " + str(light3))
print("")

#But if we find out we are running late, we may want to take more risks.
print("Oh no! I'm running late!")
print("")

light4 = stop_or_go_risky('green')
light5 = stop_or_go_risky('yellow')
light6 = stop_or_go_risky('red')

print("Should I go at a green light?... " + str(light4))
print("Should I go at a yellow light?... " + str(light5))
print("Should I go at a red light?... " + str(light6))
```

---
### Question 2: Hailstone

The Hailstone sequence from Hofstadter is as follows:
1. Pick a positive integer n as the start value.
2. If n is even, divide it by 2.
3. If n is odd, multiply it by 3 and add 1.
4. If you continue this process, n will eventually reach 1. 

Based on what you learned today, write a function that uses if/else statements to do **one step** of the hailstone sequence. It should take in an integer n and then return the corresponding Hailstone number for that value. For example, `hailstone(10)` should return 5 and `hailstone(9)` should return 28. 


```python
def hailstone(n):
    if ...:
        ...
    else:
```

```python
grader.check("q2")
```

```python
# Run this cell to see the Hailstone sequence in action!
print("Starting value: 10")
n1 = hailstone(10)
print("Step 1:", n1)
n2 = hailstone(n1)
print("Step 2:", n2)
n3 = hailstone(n2)
print("Step 3:", n3)
n4 = hailstone(n3)
print("Step 4:", n4)
n5 = hailstone(n4)
print("Step 5:", n5)
n6 = hailstone(n5)
print("Step 6:", n6)
```

---
### Question 3: Fizzbuzz

Create a function named `fizzbuzz` that follows the classic Fizzbuzz game rules for a given integer `n`. Here are the rules:


1. return `"fizz"` if `n` is divisible by 3.
2. return `"buzz"` if `n` is divisible by 5.
3. return `"fizzbuzz"` if `n` is divisible by both 3 and 5.
4. return the number itself as a string if it's neither divisible by 3 nor 5.

Implement the fizzbuzz function to meet these specifications. Your implementation should handle any positive integer n.

*Hint:* The order of your conditions matter!

*Hint:* We learned `if` and `else` in class on Monday. You should also use `elif` in your implementation. See [this resource](https://www.w3resource.com/python/python-if-else-statements.php) for more information.

```python
def fizzbuzz(n):
    if ...:
        ...
    elif ...:
        ...
    elif ...:
        ...
    else:
```

```python
grader.check("q3")
```

## Done! 😇

## Pets of Data 6
Snickers wishes you a happy rest of the week!

<img src="snickers.jpg" width="40%" alt="One dog looking at a treat"/>

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> isinstance(x_satifies, bool)
True
```
```
>>> x_satifies is True
True
```


Tests for q2 (points: 1)
```
>>> callable(hailstone)
True
```
```
>>> hailstone(2) == 1
True
```
```
>>> hailstone(3) == 10
True
```
```
>>> (hailstone(10) == 5) and (hailstone(9) == 28)
True
```


Tests for q3 (points: 1)
```
>>> callable(fizzbuzz)
True
```
```
>>> fizzbuzz(15) == 'fizzbuzz'
True
```
```
>>> fizzbuzz(3) == 'fizz'
True
```
```
>>> fizzbuzz(5) == 'buzz'
True
```
```
>>> fizzbuzz(2) == '2'
True
```
