```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Homework 3 – Advanced Table Methods and Data Visualization

## Data 6, Fall 2024

This homework is due on **Thursday, October 31st at 11:00PM**. 

```python
# Run this cell, don't change anything.
from datascience import *
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
plt.style.use("ggplot")
from ipywidgets import interact, IntSlider
import warnings
warnings.filterwarnings('ignore')
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

# Part 1: Visualization Fundamentals

<br></br>
<hr style="border: 1px solid #fdb515;" />

## Acknowledgements

Many of the pictures and descriptions in this section of the homework are taken from [Dr. Allison Horst](https://twitter.com/allison_horst) and [Dr. Kristen Gorman](https://www.uaf.edu/cfos/people/faculty/detail/kristen-gorman.php). The data was collected as a part of the [Long Term Ecological Research Network](https://lternet.edu/). If you're curious about the origins of the data, see [here](https://allisonhorst.github.io/palmerpenguins/) for more details.

<img src='lter_network.jpeg' width=300>
    
<br></br>
<hr style="border: 1px solid #fdb515;" />

## The Data

In this part of your assignment, we will explore a dataset containing size measurements for three penguin species observed on three islands in the Palmer Archipelago, Antarctica. The data was collected by Dr. Kristen Gorman, a marine biologist from 2007 to 2009.

Here's a photo of Dr. Gorman in the wild collecting the data:

<img src='gorman1.png' width=500>

Run the cell below to load in our data.

```python
# Run this cell!
penguins = Table.read_table("penguins.csv")
penguins
```

<hr style="border: 1px solid #fdb515;" />

## The Columns

Let's make sure we understand what each of the columns in our data represents before proceeding.

**`"species"`** There are three species of penguin in our dataset: Adelie, Chinstrap, and Gentoo.
<img src='lter_penguins.png' width=500>

**`"island"`** The penguins in our dataset come from three islands: Biscoe, Dream, and Torgersen. (The smaller image of Anvers Island may initially be confusing; the dark region is land and the light region is water.)

<img src='island.png' width=500>

Image taken from <a href="https://journals.plos.org/plosone/article/figure?id=10.1371/journal.pone.0090081.g001">here</a>


**`"bill_length_mm"`** and **`"bill_depth_mm"`**: See the illustration below.

<img src='culmen_depth.png' width=350>

**`"flipper_length_mm"`**: Flippers are the equivalent of wings on penguins. 

**`"body_mass_g"`** The body mass of the penguin, in grams

**`"sex"`**: The sex of the penguin

<hr style="border: 1px solid #fdb515;" />

# Question 1 - Penguins!

<!-- BEGIN QUESTION -->

---
## Question 1a
Let's start by visualizing the distribution of the islands from which the penguins in our dataset come. Your friend Sunya says that using a **histogram** is the best way to do this. However, Josh disagrees and says that using a **horizontal bar chart** is a better move. Who is right? Explain your answer.

*Hint*: What kind of variable is contained within the `"island"` column?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
## Question 1b

In order to visualize this distribution, create a table named `islands` with the following columns:
1. `island`: A column containing all the possible islands a penguin could come from
2. `num penguins`: A column containing the count of all penguins that reside in that island.

```python
islands = ...
```

```python
grader.check("q1b")
```

---
## Question 1c

Based on your answer for the previous question, create an appropriate visualization to show the distribution of the islands from which the penguins come.

Specifically, write a line of code in the following cell that produces a visualization with three bars, one for each island. The length of each bar should correspond to the number of penguins from each island.

Use your `islands` table from the previous problem.

<!-- BEGIN QUESTION -->



```python
...

# Do not change any code below this comment
plt.xlabel("Number of Penguins")
plt.title("Number of Penguins on Each Island");
```

<!-- END QUESTION -->

---
## Question 1d

Let's construct a table with information about the penguins species.

This table should be called `species`, and contains the following columns:
1. `"species"`: A column containing all of the possible `species` in the `penguins` table
2. `"bill_length_mm median"`: The median bill length of each respective `species`.
3. `"bill_depth_mm_median"`: The median bill depth of each respective `species`.
4. `"flipper_length_mm median"`: The median flipper length of each respective `species`.
5. `"body_mass_g_median"`: The median body mass of each respective `species`.


*Hint*: You may find several functions from the [Python Reference](https://data6.org/fa24/reference/) helpful

```python
species = ...
species
```

```python
grader.check("q1d")
```

---
## Question 1e

Using the `species` table from the previous question, construct a new table called `species_new` containing the following columns: 
1. `"species"`: The species column from the `species` table
2. `"bill_length_mm median"`: The bill length median column from the `species` table
3. `"bill_depth_mm_median"`: The bill depth median column from the `species` table
4. `"code"`: A new column which contains the **first letter** of each species name

Hint: Below we've defined a function `first_letter`, which takes in a String and returns the first letter of that String. DO NOT HARDCODE THE VALUES A, C, and G.

```python
def first_letter(name):
    """Returns the first letter of the name.
    Inputs: String representing the name
    Returns: String representing the first character in the name.
    """
    return name[0]
```

```python
species_new = ...
species_new
```

```python
grader.check("q1e")
```

<hr style="border: 1px solid #fdb515;" />

# Question 2 - Histograms

<img src='all3.png' width=400>


Adelie, Chinstrap, and Gentoo penguins.


Great! Now that we've explored the distributions of some of the categorical variables in our dataset (species and island), it's time to study the distributions of some of the numerical variables. As covered in Lecture 15, we can visualize a numerical distribution by creating a histogram.

---
## Question 2a

Before we produce any histograms using code, let's make sure we understand *how* to interpret the data plotted by a histogram. Run the following cell to create a histogram that shows the **distribution of penguin body mass (in grams)** for all the penguins in the `penguins` table.

```python
penguins.hist("body_mass_g", bins=np.arange(2000,7001, 500), density=False)
```

**Task**: Based on the information presented in the above histogram, assign `num_penguins` to the **number of penguins** that have a body mass **between 3000 grams (*inclusive*) and 3500 grams (*exclusive*)**.

*Note*: Remember, in a histogram, each bin is inclusive of the left endpoint and exclusive of the right endpoint. What this means is that, for example, the bin between 2500 and 3000 grams above corresponds to body masses that are greater than or equal to 2500 grams and less than 3000 grams.


```python
num_penguins = ...
num_penguins
```

```python
grader.check("q2a")
```

<!-- BEGIN QUESTION -->

---
## Question 2b

Sandra takes a quick glance at the histogram from **Question 2a** and claims the following:
> _"I can tell you how many penguins have body masses between 5000 and 5200 grams, **only** using the histogram above."_

Leanne disagrees with Sandra, stating:
> _"That's not true. You'd need more than just the histogram above to know that information."_

Who is correct in this situation: Sandra or Leanne? Explain why the *other person is incorrect*, providing justifcation based on what you've learned about histograms.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

---
## Question 2c

Now it's time to create some histograms! For convenience, we've provided the `penguins` table in the cell below. There are four columns containing numerical variables: `"bill_length_mm"`, `"bill_depth_mm"`, `"flipper_length_mm"`, and `"body_mass_g"`.

```python
penguins.show(5)
```

<!-- BEGIN QUESTION -->

**Task**: In the cell below, write code that produces a histogram that visualizes the distribution of our penguins' flipper lengths. We've provided the bins for you to use as the `my_bins` variable. More specifically, you'll want to re-create the following histogram exactly:

<img src="flipper_lengths.png">


```python
my_bins = np.arange(170,240, 5)
...

# Do not change anything below this line
plt.title("Distribution of Flipper Lengths (mm)");
```

<!-- END QUESTION -->

---
## Question 2d

When creating histograms, it's important to try several different bin sizes in order to make sure that we're satisfied with the level of detail (or lack thereof) in our histogram.

Run the code cell below. It will present you with a histogram of the distribution of our penguins' body masses, along with a slider for bin widths. Use the slider to try several different bin widths and look at the resulting histograms.

```python
# Don't worry about the code, just play with the slider that appears after running.
def draw_mass_histogram(bin_width):
    penguins.hist('body_mass_g', bins = np.arange(2700, 6300+2*bin_width, bin_width), density = False)
    
interact(draw_mass_histogram, bin_width=IntSlider(min=25, max=2000, step=25, continuous_update=False));
```

<!-- BEGIN QUESTION -->

**Task**: In the cell below, compare the two histograms that result from setting the bin width to **100** and **750.** What are the pros and cons of each size? (Remember that these histograms are displaying the same data, just with different bin sizes.)


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

---
## Question 2e

In the previous question, you plotted the distribution of penguins' body mass measurements for the **entire dataset**. While doing so is important to visualize our data as a whole, we can uncover underlying patterns within our data by grouping. This process is known as **disaggregation**: separating aggregated data into its component parts; we saw this process in Lab 2 with the BRFSS data set.

**Task**: Using the `penguins` table and the `hist` table method and its various optional arguments, re-create the following set of histograms:

<img src="mass_disaggregated.png" width=250>

*Note*: Use the `my_bins` variable we've defined as the optional `bins` arugment of the `hist` method.


```python
my_bins = np.arange(2500, 6500, 500)
...
```

<!-- END QUESTION -->

## Part 3: Real-World Risks within Data

Now, we are going to look at a very interesting dataset, the 1973's UC Berkeley Admission for data analysis! Through out this part, we will be working with `UCBerkeley1973_Admission.csv` file.

```python
cal_data  = Table.read_table("UCBerkeley1973_Admission.csv")
cal_data
```

<!-- BEGIN QUESTION -->

### Question 3a (Discussion)
Looking at `cal_data` first glance, what do you observe? Try exploring the dataset by what you have learned so far in class! Reflect in 2-3 sentences on what you have noticed about the data. Feel free to add/ delete more cells beyond our given cell.

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# You can use this cell to explore the `cal_data` table
```

For this case study, we'll be doing some exploration of the rates of admission between males and females at UC Berkeley in 1973. To start off, below we calculate the raw acceptance rates for males and females:

```python
total_f = sum(cal_data.column("Gender") == "F")
total_m = sum(cal_data.column("Gender") == "M")
accepted_f = cal_data.where("Admission", "Accepted").where("Gender", "F").num_rows
accepted_m =  cal_data.where("Admission", "Accepted").where("Gender", "M").num_rows
acceptance_rate_f = accepted_f / total_f * 100
acceptance_rate_m = accepted_m / total_m * 100
print("1973's Berkeley admission rate seems to be: female:", acceptance_rate_f, "and male:", acceptance_rate_m)
```

Keep these rates in mind as we begin exploring some visualization and table methods to look into an effect called **"Simpson's Paradox"**. 

---
### Simpson's Paradox
According to [Wikipedia](https://en.wikipedia.org/wiki/Simpson%27s_paradox), Simpson's Paradox "is a phenomenon in probability and statistics in which a trend appears in several groups of data but disappears or reverses when the groups are combined."

At first glance of the UC Berkeley Admission dataset, it seems male have higher chance being accepted. However, when looked closer by major (department), a different story emergered. We found that females were equal or more likely to be accepted when looked at each department. In fact, a [study](https://homepage.stat.uiowa.edu/~mbognar/1030/Bickel-Berkeley.pdf) found that female applicants disproportionately applied to highly competitive majors with low acceptance rates (e.g., English), while men more often applied to less competitive majors with higher acceptance rates. *More female applicants applying to competitive programs skewed the combined data, creating the appearance of bias against female applicants, when in many cases, women were either equally likely or more likely to be admitted compared to men in the same department.*

Now let's see this phenomenon in practice! 

We have manipulated the `cal_data` and got the numbers for rejected and accepted count for both female and male in our 1973 admission data.

### Question 3b
Assign the table `admissions` to one that contains the count of individuals rejected and accepted for both females and males.

*Hint:* Using the `cal_data` table, we want the `"Gender"` column to be our columns, and the `"Admission"` column to be our rows.

```python
admissions = ...
admissions
```

```python
grader.check("q3b")
```

### Question 3c
We want to manually add a new column called `"Acceptance Rate"` with values of acceptance rates for females and males to the `admissions` table. Below, fill in the code to calculate the female and male acceptance rates, and from there, we supply the code that creates a new column with these rates. Acceptance rate should be calculated as accepted count divided by total count!

*Hint:* You may find it helpful to use the `.column()` and `.item()` methods.

```python
f_acceptance_rate = ...
m_acceptance_rate = ...
acceptance_rates = make_array(f_acceptance_rate, m_acceptance_rate)
admissions = admissions.with_column("Acceptance Rate",  acceptance_rates)
admissions
```

```python
grader.check("q3c")
```

<!-- BEGIN QUESTION -->

### Question 3d
Using the adjusted `admissions`, create a bar graph comparing the acceptance rate between female and male applicants.

```python
...
```

<!-- END QUESTION -->

### Question 3e
Which of the following most accurately describes the takeaway from the visualization created by `admissions` above? You should answer the question by assigning `q3e` to `make_array(...)` where `...` is the choice of your answer (eg. `make_array(3)`).

1. Female applicants have higher chance of being accepted.
2. Male applicants have higher chance of being accepted.
3. Female and male applicants have the same amount chance of being accepted.

```python
q3e = ...
```

```python
grader.check("q3e")
```

We see by simply utilizing `"Gender"` and `"Admission"` to calculate overall acceptance rates for graduate admissions, the male acceptance rate seems to be higher than the female acceptance rate. However, besides `"Gender"` and `"Admission"`, there was a third column in this table: `"Major"`. Let's investigate how the major will take into play with admission rate between these two recorded gender. (How "grouping" by major as described in Simpson's Paradox will affect the applicants' acceptance rate by gender).

Below, we create two different tables using more complex functionalities of `.group()` and `.pivot()`. Take a look at the information we have gained using the `"Major"` feature below:

```python
admission_major = cal_data.pivot('Gender', 'Major', collect = lambda x: sum(x == 'Accepted') / len(x), values = 'Admission')
admission_major = admission_major.relabeled(['F', 'M'], ['F Acceptance Rate', 'M Acceptance Rate'])
admission_major
```

```python
num_applicants = cal_data.group(['Major', 'Gender']).pivot('Gender', 'Major', collect = np.sum, values = 'count')
num_applicants = num_applicants.relabeled(['F', 'M'], ['F Application Count', 'M Application Count'])
num_applicants
```

<!-- BEGIN QUESTION -->

### Question 3f
Using `admission_major` and `num_applicants`, create two separate overlaid bar charts. The first one will plot the number of applicants of males and females across the six majors, and the second will plot their acceptance rates in the majors.

```python
...
...
```

<!-- END QUESTION -->

### Question 3g
Which of the following multiple choice correcrly describes *the visualization created above* with `admission_major` and `num_applicants`? You should answer the question by assigning `q3g` to `make_array(...)` where `...` is the choice of your answer (eg. `make_array(3, 4)`). Your choice can be can multiple.

1. Among the 6 different majors (not including `"Other"`, female applicants have a slightly higher chance of being accepted than male applicants in several majors
2. Among the 6 different majors (not including `"Other"`), male applicants have a slightly higher chance of being accepted than female applicants in several majors.
3. Overall, female applicants has higher chance of getting in to UC Berkeley.
4. Overall, male applicants has higher chance of getting in to UC Berkeley. 
5. There seems to be **higher female applicants** to the major when there are **lower female acceptance rate** compared to male acceptance rate.  
6. With the visualization above, we can claim there is structrual inequity against female women applicants. 

```python
q3g = ...
```

```python
grader.check("q3g")
```

---
### Thinking about Implications 🤔

<!-- BEGIN QUESTION -->

### Question 3h (Discussion)
Consider the takeaways that you had after looking at the different visualizations you created in **Questions 3d and 3f**. Based on what we've discussed regarding the various sociological paradigms, how do you think these conclusions best align with the paradigms?

*Hint:* The conclusions might not line up perfectly with any one of the paradigms! Just feel free to reflect on which ones they might line up best or worst with. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 3i (Discussion) 
A media company wants to make a report on UC Berkeley admissions. What would happen if they ONLY report with the visualization you created in **`Question 3d`**? How would this report potentially impact the society we see today?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# Done!

Congrats! You've finished another Data 6 homework assignment! To submit your work, submit the zip file to the autograder assignment and the pdf to the written assignment, both in gradescope. Make sure to assign your pages for the written assignment!

This homework is out of **38 points**. The point breakdown for this assignment is given in the table below:

| **Category** | Points |
| --- | --- |
| Autograder | 18 |
| Written | 20 |
| **Total** | 38 |

## Pets of Data 6
**Spark** hopes you enjoyed Homework 2!

<img src="spark.jpg" width="40%" alt="Brown dog chillin at the beach."/>

## Submission

Below, you will see two cells. Running the first cell will automatically generate a PDF of all questions that need to be manually graded, and running the second cell will automatically generate a zip with your autograded answers. You are responsible for submitting both the coding portion (the zip) and the written portion (the PDF) to their respective Gradescope portals. **Please save before exporting!**

> **Important: You must correctly assign the pages of your PDF after you submit to the correct gradescope assignment. If your pages are not correctly assigned and/or not in the correct PDF format by the deadline, we reserve the right to award no points for your written work.**

If there are issues with automatically generating the PDF in the first cell, you can try downloading the notebook as a PDF by colicking on `File -> Save and Export Notebook As... -> PDF`. If that doesn't work either, you can manually take screenshots of your answers to the manually graded questions and submit those. Either way, **you are responsible for ensuring your submision follows our requirements, we will NOT be granting regrade requests for submissions that don't follow instructions.**

```python
from otter.export import export_notebook
from os import path
from IPython.display import display, HTML
name = 'hw03'
export_notebook(f"{name}.ipynb", filtering=True, pagebreaks=True)
if(path.exists(f'{name}.pdf')):
    display(HTML(f"Download your PDF <a href='{name}.pdf' download>here</a>."))
else:
    print("\n Pdf generation failed, please try the other methods described above")
```

```python

```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1b (points: 3)
```
>>> isinstance(islands, Table)
True
```
```
>>> islands.labels == ('island', 'num penguins')
True
```
```
>>> islands.num_rows == 3
True
```
```
>>> islands.where("island", "Biscoe").column('num penguins').item(0) == 163
True
```


Tests for q1d (points: 3)
```
>>> isinstance(species, Table)
True
```
```
>>> species.num_rows == 3
True
```
```
>>> species.labels == ("species", "bill_length_mm median", "bill_depth_mm median", "flipper_length_mm median", "body_mass_g median")
True
```
```
>>> species.where("species", "Adelie").column("bill_depth_mm median").item(0) == 18.4
True
```


Tests for q1e (points: 3)
```
>>> isinstance(species_new, Table)
True
```
```
>>> species_new.labels == ('species', 'bill_length_mm median', 'bill_depth_mm median', 'code')
True
```
```
>>> species_new.num_rows == 3
True
```
```
>>> species_new.where("species", "Gentoo").first("code") == "G"
True
```


Tests for q2a (points: 1)
```
>>> any(isinstance(num_penguins, x) for x in (int, float))
True
```
```
>>> 58 <= num_penguins <= 62
True
```


Tests for q3b (points: 2)
```
>>> isinstance(admissions, Table)
True
```
```
>>> admissions.column("Accepted").item(0) == 1494
True
```


Tests for q3c (points: 3)
```
>>> isinstance(admissions, Table)
True
```
```
>>> round(admissions.column("Acceptance Rate").item(0), 6) == 0.345753
True
```


Tests for q3e (points: [0, 1])
```
>>> isinstance(q3e, np.ndarray)
True
```
```
>>> q3e.item(0) == 2
True
```


Tests for q3g (points: [0, 0.5, 0.75, 0.75])
```
>>> isinstance(q3g, np.ndarray)
True
```
```
>>> len(q3g) == 2
True
```
```
>>> 1 in q3g
True
```
```
>>> 5 in q3g
True
```
