```python
# Initialize Otter
import otter
grader = otter.Notebook("lab09.ipynb")
```

<img src="data6.png" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

# Lab 9  - Iteration

## Data 6, Fall 2024

Welcome back to Lab 9! This week we will be covering content from lectures on **control**, **iteration**, and **strings**.

```python
# Just run this cell
from datascience import *
import numpy as np
```

# Part 1: Control

When writing functions in Python, we may want the function to behave differently depending on the input. We could choose to write several similar functions to accomplish this, but that would require copying and pasting much of the same code over and over again, only making small changes.

Instead, we need a way to tell one singular function to execute different code for different inputs. In lecture we discussed this idea as `if-else` statements.

Let's look at an example of a function that uses `if-else` statements to tell us what water looks like at different temperatures:

<img src='water.jpeg' width=300>

```python
def state_of_water(temperature):
    if temperature <= 32: # Water is solid at and below 32°F
        return "ice"
    elif temperature < 212: # Water is liquid between 32°F and 212°F
        return "liquid"
    else: # Water is gaseous above 212°F
        return "steam"
```

Here, we could have made 3 different functions that each dealt with a different state of water, but that would force us to know the state of water (what we currently do not know and are trying to figure out) so that we could pick the right function to use. This is why control such as `if-else` statements are so important!

Now let's use this function to determine the state of water in the following locations during the winter:

| **City** | Temperature (°F) |
| --- | --- |
| Berkeley | 58 |
| New York | 24 |
| Miami | 78 |
| Earth's Core | 10800 |

We can see the calculation each city's state of water using the function defined above:

```python
print(f"Water in Berkeley is: {state_of_water(58)}")
```

Output:
```
Water in Berkeley is: liquid
```

```python
print(f"Water in New York is: {state_of_water(24)}")
```

Output:
```
Water in New York is: ice
```

```python
print(f"Water in Miami is: {state_of_water(78)}")
```

Output:
```
Water in Miami is: liquid
```

```python
print(f"Water at the Center of the Earth is: {state_of_water(10800)}")
```

Output:
```
Water at the Center of the Earth is: steam
```

We can call the function on different inputs and the function deals with each input differently based on the control logic you gave it. This will be very important when you write functions that deal with large, sometimes unpredictable datasets where the logic of your function will deal with inputs you may not have directly prepared the function for.

# Part 2: Iteration

Iteration is the repetition of code. Typically, the variables in iteration change sequentially, and the iteration continues to repeat until a boolean condition indicates that the iteration be stopped. 

---
## Question 1

A `for` loop is an example of iteration in Python that we have seen in lecture. Let's use a `for` loop to write a function that finds the sum of the digits of a number.

As an example, the sum of the digits of the number 123 is 1 + 2 + 3 = 6, so `sum_digits(123)` should return 6.

```python
def sum_digits(n):
    """
    Sums all the digits of n.
    >>> sum_digits(123)
    6
    >>> sum_digits(1)
    1
    >>> sum_digits(701)
    8
    """
    counter = 0
    str_n = str(n)
    i = 0
    while i < len(str_n):
        counter = counter + int(str_n[i])
        i = i + 1
    return counter
sum_digits(701)
```

Output:
```
8
```

```python
grader.check("q1")
```

# Part 3: Strings

We have already seen strings before in this course, but we want to introduce some new ideas with strings that will be very useful when dealing with data from real-world data sets.

Strings can be viewed as special lists of characters. They can be printed, sliced, indexed, and more! You can do all of these things with lists as well, so let's take a look at how strings and lists are similar:

```python
our_string = "Data Science at Berkeley rocks!"
print(our_string)
print(our_string[0] + our_string[5] + "@" + our_string[16:24])
```

Output:
```
Data Science at Berkeley rocks!
DS@Berkeley
```

We can iterate and slice strings too! In this way, strings are like special arrays.

There is an important fact to point out here: In Python,  `strings` are **immutable**. This means that you are not able to mutate (i.e. change) a string object in place. If you want to modify your string, you'll have to make an entirely new one.

```python
our_string
```

Output:
```
'Data Science at Berkeley rocks!'
```

```python
# However you cannot do this with strings because they are immutable.
# This will be important when we talk about cleaning data sets before we can use them.
our_string[3] = 'x'
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[16], line 3
      1 # However you cannot do this with strings because they are immutable.
      2 # This will be important when we talk about cleaning data sets before we can use them.
----> 3 our_string[3] = 'x'

TypeError: 'str' object does not support item assignment
```

There are special string methods that help us edit strings, but we cannot forcefully change them with indexing or slicing. If you want to change a letter or a part of a string, you need to use those `str` methods from lecture.

---
## Question 2

Let's use this idea of array and string indexing and `if/else` to find all words that begin with a vowel. Complete the function `starts_with_vowels` so that it takes an array of strings, finds all of the words that begin with a vowel, and then returns an array of said words.

For the example array, `some_words`, your function should return `array(['alligator', 'aardvark', 'owl', 'unicorn'])`.

*Hint*: You'll need to use the `np.append` function, the documentation for which you can find [here]("https://numpy.org/doc/stable/reference/generated/numpy.append.html").

```python
some_words = make_array("alligator", "monkey", "zebra", "lion", "aardvark", "owl", "bear", "unicorn")
another_list_of_words = make_array("rocco", "ted", "sophia")
def starts_with_vowels(words):
    words_to_keep = make_array()
    for str in words:
        if str[0] in "aeiou":
            words_to_keep = np.append(words_to_keep, str)
    return words_to_keep
starts_with_vowels(another_list_of_words)
```

Output:
```
array([], dtype=float64)
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

---
## Question 3

Let's write a function that uses another `for` loop. We want to be able to find the acronym of a name. For example, we want to be able to call **The University of California, Berkeley** by an abbreviated `UCB`.

Because we don't care as much about small words even if they are capitalized, also make sure to **exclude words with three letters or fewer from your acronym**. That is why the `T` in `The` is not included in the example acronym above.

Write a function that takes in a string and returns its acronym.

_Hint_: the string method `.isupper()` will return `True` if the given string (or character) is entirely upper case, and `False` otherwise.

```python
# Here is a cell that explains the str.split() function used in the solution below:
sentence = "Hello, I am a student in Data 6!"
sentence.split()
```

```python
def acronym(title):
    # This str.split() function mentioned above turns a string into a list of its words.
    words = title.split()
    result = ""
    ...
    return result

acronym("The University of Califiornia, Berkeley")
```

```python
grader.check("q3")
```

## Done! 😇

This lab is optional, but you're welcome to submit to the autograder on gradescope if you'd like. 

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1 (points: 1)
```
>>> callable(sum_digits)
True
```
```
>>> sum_digits(123) == 6
True
```
```
>>> sum_digits(0) == 0
True
```
```
>>> sum_digits(1234) == 10
True
```


Tests for q2 (points: 1)
```
>>> all(starts_with_vowels(some_words) == make_array('alligator', 'aardvark', 'owl', 'unicorn'))
True
```
```
>>> all(starts_with_vowels(make_array("bear")) == make_array())
True
```
```
>>> "zebra" not in starts_with_vowels(some_words)
True
```


Tests for q3 (points: 1)
```
>>> callable(acronym)
True
```
```
>>> acronym("The University of California, Berkeley") == "UCB"
True
```
```
>>> acronym("") == ""
True
```
```
>>> acronym("Data Science at Berkeley") == "DSB"
True
```
```
>>> acronym("The United States of America") == "USA"
True
```
